# Exercise 2

## physics718: Programming in Physics and Astronomy with C++ or Python (SS 2020)
Oliver Cordes & Thomas Erben

Homework is due on **Tuesday, 05/05/2020, 23:59pm**

 * You only learn a programming language by actively praticing and using it! We therefore **strongly** advise you to work on the homework problems.
 * Please discuss the problems with your student peers in case of problems.
 * Your code(s) needs to be well and appropriately commented!
 * All files of your solution should be in the same directory as this notebook
 * Finally submit this notebook with your solutions in the nbgrader system
 
**Main topics of this exercise:**
 * work with scalar types in Python *int*, *float* and *bool*
 * control structures *if* and *while*
 * floating poing accuracy

Put in here your group number: 

Group 12

# 1. Numbers to binaries (20 points)

## 1.1) Integers to binaries

As mentioned in the lecture integer numbers can be written in different number systems. For us humans we are working in the decimal system but the computers are dealing with numbers in the binary system. Theoretically all integer numbers $n$ can be written in number systems with a base $b$:

$$ n = \sum_{i=0}^{m} q_{i}b^{i}$$

with $0 \le q_{i} < b$. As said we humans use $b=10$ and the computer $b=2$. 

In this first task we want to convert numbers into the binary system. Here are some examples:
* $1_{10}$ -> $1_2$
* $3_{10}$ -> $11_2$
* $14_{10}$ -> $1110_2$
etc.

For the conversion you can use the following algorithm. Start with the number to be converted and divide this number by 2 and store the quotient and the rest. Then proceed with the quotient and divide this number again until the quotient reaches 0. The rest numbers written in the reverse order gives the representation of $n$ in the binary system.

Here is an example which converts $n=14$ into the binary system:
$$ 14 :2 = 7 + 0$$ 
the quotient is 7  and the rest is 0, then:
$$ 7 :2 = 3 + 1 $$
$$ 3 :2 = 1 + 1 $$
$$ 1 :2 = 0 + 1 $$

The rest numbers written in the reverse order are now $1110$ which is 14 written in the binary system!

**Your task:**

Implement the algorithm in a single notebook code cell. Start with a predefined variable $n$ which contains the number to convert. The converted number should be stored in a string with the name *output_integer*. Print this string at the end of the code. The value of *output_integer* will be tested in the grading system, so for the validation please convert the $n=14$!

**Hints:**
 * on this [webpage](https://indepth.dev/the-simple-math-behind-decimal-binary-conversion-algorithms/) you can see how and why the algorithm is working (this page uses the base 2 and describes the decimal to binary conversion)
 * you can do all the divisions in a while loop
 * for calculating the quotient and the rest you can use the [*divmod*](https://docs.python.org/3/library/functions.html#divmod)-command in python.
   
   Example:
   ```Python
   quot, rest = divmod(a,b)
   ```
 * for the representation of the converted number you can start with an empty string e.g. ```s = ''``` and prepend a rest
   number in every step with ```s = str(rest) + s``` !
 * test your code with different examples, e.g. $n=14$ and $n=0$!

<!-- grade_id: cell-249cdf3b0452f1ac, points: 0, grade: false, solution: true -->
```python
#Choose an integer n
n = 14
b = 2

output_integer = ''
quot = n

#Check whether the input number is zero.
if (quot == 0):
    output_integer = str(0)
    
#If the input number is not zero, at each step divide it by 2 and 
#store the remainder in a string. The cycle stops when the quotient is zero.
else:
    while (quot > 0):
        quot, rest = divmod(quot,b)
        output_integer = str(rest) + output_integer
        
#Print the final string.       
print(output_integer)
```

<!-- grade_id: cell-bcf437cbac810793, points: 10, grade: true, solution: false -->
```python
assert(output_integer == '1110')
```

## 1.2) Floats to binaries

Converting floating point number into the binary system is similar to the integer conversion. The sum as shown above can be extended with negative indices. 

$$ x = \sum_{i=-k}^{m} q_{i}b^{i}$$

Here are a few examples:
 * $1.5_{10}$  -> $1.1_2$
 * $3.25_{10}$ -> $11.01_2$
 ...
 
 
For the conversion of a float number $x$ one can use the following algorithm. First we will multiply $x$ with 2 until $x$ is (approximately) an integer. We should store the number of multiplications. Then we can convert this int into a binary using the algorithm which you implemented above. To convert the binary into a binary float we can add a *binary point* at the end and shift the point to the left by the number of multiplication positions.

Let's convert the floating point number $x=3.25$:

1. Do the multiplications: $3.25 * 2 = 6.5$ -> $6.5 * 2 = 13$ -> which means 2 multiplications until we have an integer
2. convert 13 into a binary -> $1101$
3. add a binary point at the end: $1101.$
3. shift the point by 2 positions to the left, to compensate for the multiplications -> $x_b = 11.01_2$

A second example, convert $x=0.125$:

1. Do the multiplications: $0.125 * 2 * 2 * 2 = 1$ -> 3 multiplications
2. convert 1 into a binary (trivial) -> $1$
3. now we need to prepend $n*0$ (n zeros, n = min. the number of multiplications) and the binary point at the end: $0001.$
4. shift the point by 3 positions to the left, to compensate for the multiplications -> $x_b = 0.001_2$

**Your task**:

Implement this algorithm. You can use the code from the previous task to convert an integer into a binary. 
For some of the floating point numbers there will be not an exact binary representation. You can stop the multiplication after you calculated a value which differs less than $10^{-6}$ from its integer representation. 
If you have done more than 10 multiplications you should print a warning at the end of your code indicating that the binary representation is probably not finite.

**Hints**:
 * to convert an floating point number $x$ to an integer you can use ```int(x)```! ```int(x)``` truncates $x$ towards zero.
 * use the function ```abs(x-int(x))``` to test the accuracy of the interger conversion
 * to shift the binary point, you can directly put the binary point at the correct place (don't add the point at the end before).    For adding the point, please  have a look at the following cell with a small introduction into *slices* of array and array like structures e.g. strings

```python
# strings are represented similar to arrays
s = '123456789'

# each character can accessed by their position
print(s[0])   # the first character, starting with 0 until length of the string -1
print(s[-1])  # negative indices counts from the end

# you can cut a smaller string out of the larger string, by defining start and end-indices
print(s[1:3]) # print the second and the third character
print(s[1:])  # if you leave the second limit open, it means until the end
print(s[:2])  # if you leave the beginning open, it means from the beginning
print(s[-2:]) # print the last two characters 

point = 5
# if you want to place a point between the the first  and the last 5 characters
# you need to cut the string into 2 pieces and put the '.' in between
s = s[:-shifts] + '.' + s[-shifts:]
print(s)
```

<!-- grade_id: cell-95b6c4063d5857c1, points: 0, grade: false, solution: true -->
```python
#This code calculates the binary representation of a float number in decimal representation.

import warnings

#Must be a float 
x = 3.25

b = 2
m = 0
output_integer = ''
output_float = ''

#It multiplies the number x by 2. The cycle stops when x is an integer
#or when it finds an integer such that the distance between x and the integer
#is less than 10^-6.
while ( abs(x-int(x)) > 1e-06 ):
    x = x*2
    m = m + 1 #It stores the number of multiplications
    #If the number of multiplications is greater than 10, prints a warning and ends the cycle.
    if (m > 10):
        warnings.warn('The binary representation is not finite!')
        break

#If no warning has been produced, it proceeds to convert x into a binary. 
#The binary is stored in output_integer.
if (m <= 10):
    quot = int(x)
    if (quot == 0):
        output_float = str(0)
    else:
        while (quot > 0):
            quot, rest = divmod(quot,b)
            output_integer = str(rest) + output_integer
        
        #It checks whether the length of output_integer is less than or equal to the number of multiplications.
        #If this is the case, it appends some zeros in front of output_integer.
        l = len(output_integer)
        if (l <= m):
            k = m - l + 1
            output_integer = k*str(0)+output_integer
            
    #It inserts a point in the string that contains the binary and shifts it
    #by m positions to the left.
    output_float = output_integer[:-m] + '.' + output_integer[-m:]
    #It prints the final string.
    print(output_float)
    

```

<!-- grade_id: cell-ab47346348a927a3, points: 10, grade: true, solution: false -->
```python
assert(output_float.lstrip('0') == '11.01')
```

# 2. Problems with an integral series (15 points)

Consider the sequence
$$
E_n=\int_0^1x^n\exp(x-1)\,\mathrm{d}x; \qquad n=1,2,\dots
$$
One can show that it converges to zero: $\lim_{n\to\infty}E_n=0$. Integrating by parts, we obtain the following relation between two consecutive elements:
$$
E_n = 1-nE_{n-1} \text{ with } E_1=1-\int_0^1\exp(x-1)\,\mathrm{d}x=\exp(-1).
$$
We want to estimate $E_{20}$. Do this by computing and printing the first 20 elements of the series. Obtain a second estimate of $E_{20}$ by reverting the relation:
$$
  E_{n-1} = \frac 1n (1-E_{n}) \text{ with } E_{50} = 0
$$
and printing the elements $E_{50}, E_{49}, \dots, E_{20}$.

**Your tasks:**

Create a cell for each experiment/program and a markdown cell in which you argue which one of the results you trust more and try to explain why they disagree that much.


**Hints:** 
 * Assume that in the first case, $E_1$ is represented internally as a float number with an error, i.e. $E_1 = E^t_1 + \epsilon$, where $E^t_1$ is the true value and $\epsilon$ the error. We know that $\epsilon\approx 10^{-18}$ for `Python` float numbers. What happens with $\epsilon$ when you calculate new elements of the series? 
 * for the exponential function you can use the numpy module with ```import numpy``` and use the defined function ```numpy.exp(x)``` to calculate $e^x$!

<!-- grade_id: cell-3eae773e632eab70, points: 5, grade: true, solution: true -->
```python
#This code computes the first twenty elements of the series
#by using the recursion relation.

import numpy as np

result = np.exp(-1)
n = 2
print(result)

#It implements the algorithm. The results are printed in increasing order,
#i.e from E_1 to E_20.
while (n <= 20):
    result = 1 - n*result
    print(result)
    n = n + 1



```

<!-- grade_id: cell-8ec1cc6f8ddca053, points: 5, grade: true, solution: true -->
```python
#The second method for computing E_20 consists in inverting the recursion relation. 

result = 0
n = 50

#It implements the recursion relation. The results are printed in decreasing order,
#i.e from E_50 to E_20
while (n >= 21 ):
    result = (1/n)*(1 - result)
    print(result)
    n = n-1
    

```

<!-- grade_id: cell-9e84b36eaf2530ef, points: 5, grade: true, solution: true -->
The first method yields a strange value for $E_{20}$. We are led to trust the second method